---
title: "Tutorial for SticsRpacks"
output: 
  learnr::tutorial:
    theme: cerulean
    ace_theme: pastel_on_dark
    anchor_sections: FALSE
    progressive: true
description: "A tutorial to learn SticsRpacks"
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(SticsRPacks)
# library(gradethis)
knitr::opts_chunk$set(echo = FALSE)
tutorial_options(exercise.lines = 5)
# gradethis::gradethis_setup()

# Add checking JavaStics path

parameters = list(
  gen_help = TRUE,
  javastics = "JavaSTICS-1.41-stics-9.2",
  workspace_xml = "JavaSTICS-1.41-stics-9.2/example",
  workspace_txt = "example",
  use_default = TRUE,
  work_dir = file.path(normalizePath(dirname(tempdir()), winslash = "/"), "SticsRpacks")
)

work_dir <- parameters$work_dir
work_dir <- gsub(pattern = "\\\\", replacement="/", x=work_dir)

# Creating work_dir if needed
if (!dir.exists(work_dir)) dir.create(work_dir)

javastics <- file.path(work_dir, parameters$javastics, fsep ="/")

if (!dir.exists(javastics)) dir.create(javastics)

if(!any(grepl("JavaStics.exe",list.files(javastics)))) {
  zip_path <- file.path(work_dir,paste0(parameters$javastics,".zip"))
  system(paste0("curl -u stics_user:w10lptr6405 https://w3.avignon.inra.fr/forge/attachments/download/2482/",parameters$javastics,".zip --output ", zip_path))
  unzip(zip_path, exdir = javastics)
}

workspace_xml <- file.path(work_dir, parameters$workspace_xml, fsep ="/")
workspace_txt <- file.path(work_dir, parameters$workspace_txt, fsep ="/")

# Generating html help files, one for each function
if(parameters$gen_help) {
  help_dir <- file.path(work_dir, "SticsRFiles_help")
  if (!dir.exists(help_dir)) dir.create(help_dir)
  SticsRFiles:::static_help(pkg = "SticsRFiles", 
                            overwrite = FALSE, 
                            out_dir = help_dir)
}
```

## Introduction

**Décrire ici en 2 mots ce que font les paquets**

<div class="alert alert-danger">
  Test danger 2
</div>

<div class="alert alert-warning">
  Test warning
</div>

<div class="alert alert-info">
  Test info
</div>


Detailed documentation on the different packages can be found at:

* [https://sticsrpacks.github.io/SticsRFiles/]
* [https://sticsrpacks.github.io/SticsOnR/]
* [https://sticsrpacks.github.io/CroptimizR/]
* [https://sticsrpacks.github.io/CroPlotR/]

**Expliquer aussi comment trouver de l'aide sur les fonctions: ? + reference, ...**


## A foretaste of SticsRpacks 

### TP Stics


### A simple example

vignette CroptimizR ou encore plus simple (1 seul param) pour réduire le temps d’exécution …


## Finding names of STICS variables

Names of STICS variables can be retrieved from partial name or keywords using the `get_var_info` function of the [SticsRFiles package](https://sticsrpacks.github.io/SticsRFiles/index.html).

Here's its description as given in the R help panel (`? get_var_info`) and the [SticsRFiles package website](https://sticsrpacks.github.io/SticsRFiles/reference/get_var_info.html):

<blockquote>

```{r, results="asis"}
lines <- SticsRFiles:::get_from_help(file.path(help_dir,"get_var_info.html"))
cat( paste( lines , collapse="\n" ) )
```

</blockquote>


#### Exercises

* Find all STICS variables characterizing roots

```{r get_var_info1, exercise=TRUE}
# Insert your R code here!
```

```{r get_var_info1-hint-1}
# Use argument keyword="roots"! 
```

```{r get_var_info1-solution}
# SOLUTION:
get_var_info(keyword = "roots")
```


## Exploring the content of STICS files 

We're going to see in this section how to retrieve information from STICS input files.
 
### Get usm names

The function `get_usms_list` from the [SticsRFiles package](https://sticsrpacks.github.io/SticsRFiles/index.html) returns the list of USMs included in a usms.xml file.

Here's its description as given in the R help panel (`? get_usms_list`) and on the [SticsRFiles package website](https://sticsrpacks.github.io/SticsRFiles/reference/get_usms_list.html):

<blockquote>

```{r, results="asis"}
lines <- SticsRFiles:::get_from_help(file.path(help_dir,"get_usms_list.html"))
cat( paste( lines , collapse="\n" ) )
```

</blockquote>

#### Exercises

* Get the list of USMs included in the Example workspace of the JavaStics distribution, which path is stored in the R object `workspace_xml`.

```{r get_usms_list1, exercise=TRUE}
# Insert your R code here!
```

```{r get_usms_list1-hint-1}
# Use file.path function to build full path to a file from directory names and filenames: file.path(path,filename) 
```

```{r get_usms_list1-solution}
# SOLUTION:
get_usms_list(usm_path = file.path(workspace_xml, "usms.xml"))
```

<!-- ```{r get_usms_list1-check} -->
<!-- grade_result( -->
<!--   pass_if(~identical(.result, get_usms_list(usm_path = file.path(workspace_xml, "usms.xml")))) -->
<!-- ) -->
<!-- ``` -->

* Get the list of cover-crops USMs (their name starts with `"cc_"`).

```{r get_usms_list2, exercise=TRUE}
# Insert your R code here!
```

```{r get_usms_list2-hint-1}
# The argument name is particularly useful in this case ...
```

```{r get_usms_list2-solution}
#SOLUTION:
get_usms_list(usm_path = file.path(workspace_xml, "usms.xml"), name = "cc_")
```

### Get parameters values from STICS input files

Parameters values included in STICS input files can be retrieved using the function `get_param_xml` from the [SticsRFiles package](https://sticsrpacks.github.io/SticsRFiles/index.html):

<blockquote>

```{r, results="asis"}
lines <- SticsRFiles:::get_from_help(file.path(help_dir,"get_param_xml.html"))
cat( paste( lines , collapse="\n" ) )
```

</blockquote>

#### Exercises

* Get the values of clay content after decarbonation (argi) for all soils defined in the Example workspace of the JavaStics distribution.

```{r get_param_xml1, exercise=TRUE}
# Insert your R code here!
```

```{r get_param_xml1-hint-1}
# Soils are defined in the "sols.xml" file ...
```

```{r get_param_xml1-solution}
#SOLUTION:
get_param_xml(file.path(workspace_xml,"sols.xml"), c("argi"))
```

```{r get_param_xml2-setup}
result <- get_param_xml(file.path(workspace_xml,"sols.xml"), c("argi"))
```

* Plot an histogram of these values that we stored in an R object called `result`

```{r get_param_xml2, exercise=TRUE}
# Insert your R code here!
```

```{r get_param_xml2-hint-1}
# Use hist function to plot an histogram
```

```{r get_param_xml2-hint-2}
# Be careful, result is a list: argi is stored in result$sols.xml$argi
```

```{r get_param_xml2-solution}
#SOLUTION:
hist(result$sols.xml$argi)
```

* Compute the Available Water Capacity of soil `"solmais"` in mm

```{r get_param_xml3, exercise=TRUE}
# Insert your R code here!
```

```{r get_param_xml3-hint-1}
# Equation of Available Water Capacity in mm from STICS inputs is: AWC=sum((HCCF-HMINF)*epc*DAF/10)
```

```{r get_param_xml3-hint-2}
# Use select and value arguments of get_param_xml to extract information for a given soil
```

```{r get_param_xml3-solution}
# SOLUTION:
result <- get_param_xml(file.path(workspace_xml,"sols.xml"), c("HCCF","HMINF","epc","DAF"),select="sol",value="solmais")
sum((result$sols.xml$HCCF-result$sols.xml$HMINF)*result$sols.xml$epc*result$sols.xml$DAF/10)
```

### Get weather variables from STICS input files

Weather variables defined in STICS input weather files can be retrieved using the function `get_climate_txt` from the [SticsRFiles package](https://sticsrpacks.github.io/SticsRFiles/index.html):

<blockquote>

```{r, results="asis"}
lines <- SticsRFiles:::get_from_help(file.path(help_dir,"get_climate_txt.html"))
cat( paste( lines , collapse="\n" ) )
```

</blockquote>

#### Exercises

* Compute the total amount of precipitation that has fallen in Auzeville in 2013 from the file `"Auzeville.2013"` of the Javastics Example workspace.

```{r get_climate_txt1, exercise=TRUE}
# Insert your R code here!
```

```{r get_climate_txt1-hint-1}
# Precipitation is stored in the column named ttrr 
```

```{r get_climate_txt1-solution}
#SOLUTION:
result <- get_climate_txt(dirpath=workspace_xml,filename="Auzeville.2013")
sum(result$ttrr)
```

* Compute the amount of precipitation that has fallen between dates "2013-03-21" and "2013-09-20".

```{r get_climate_txt2, exercise=TRUE}
# Insert your R code here!
```

```{r get_climate_txt2-hint-1}
# Use dplyr::filter(result,Date>"2013-03-21", Date<"2013-09-20") to filter the results wrt the Dates.
```

```{r get_climate_txt2-solution}
#SOLUTION:
result <- get_climate_txt(dirpath=workspace_xml,filename="Auzeville.2013")
filtered_result <- dplyr::filter(result,Date>"2013-03-21", Date<"2013-09-20")
sum(filtered_result$ttrr)
```


## Running simulations from R

Let see now how to run STICS simulations directly from R using the function `run_javaStics` from the [SticsOnR package](https://sticsrpacks.github.io/SticsOnR/index.html).

From now on we will no longer display function help in the tutorial. Use `? function_name` or explore the packages websites (in particular the Reference tab) to display it.

### Exercises

* Run all the USMs defined in the Example workspace of the JavaStics distribution. Path to JavaStics is stored in the `javastics` R object.

```{r run_javaStics1, exercise=TRUE}
# Insert your R code here!
```


```{r run_javaStics1-hint-1}
# No need to use the usms_list argument!
```

```{r run_javaStics1-solution}
# SOLUTION:
run_javastics(javastics_path=javastics,  workspace_path = workspace_xml)
```

* Run only the USMs banana and Turmeric.

```{r run_javaStics2, exercise=TRUE}
# Insert your R code here!
```


```{r run_javaStics2-hint-1}
# Use the usms_list argument!
```

```{r run_javaStics2-solution}
# SOLUTION:
run_javastics(javastics_path=javastics,  workspace_path = workspace_xml, usms_list = c("banana","Turmeric"))
```

<div class="alert alert-info">
  See also: `gen_varmod` function to modify var.mod file in case STICS does not simulate the variables you are interested in.
</div>


## Exploring and Evaluating the Stics results
PATRICE

### Get simulations results

`get_sim`

```{r, results="asis"}
lines <- SticsRFiles:::get_from_help(file.path(help_dir,"get_sim.html"))
cat( paste( lines , collapse="\n" ) )
```

#### Exercises

> Get simulations outputs of USMs from the JavaStics distribution `example` directory

<!-- Récupérer tous les résultats simulés -->
```{r get_sim1, exercise=TRUE}
# Get all outputs
```
```{r get_sim1-hint-1}
# Path of the results is stored in workspace_xml
```

```{r get_sim1-solution}
get_sim(workspace = workspace_xml)
```

<!-- Récupérer tous les résultats simulés pour les les USMs XXX -->

```{r get_sim2, exercise=TRUE}
# Get outputs for USMs named : SugarBeet, wheat, maize, ...
```

```{r get_sim2-hint-1}
# The argument usm_name is particularly useful in this case ...
```

```{r get_sim2-solution}
#SOLUTION:
get_sim(workspace = workspace_xml, usm_name = c("SugarBeet, wheat"))
```

Récupérer les résultats pour la variable XXX et les USMs XXX et XXX

	plot stics results
		Tracer les dynamiques des variables XXX et XXX indépendamment ou superposées
	
	get_obs
  (mettre une note sur l'utilisation de usms.xml ?)

	Evaluate
		dynamiques avec des obs
		scatter_plots basiques
		scatter_plots avec gestion des symboles : demander de mettre des symboles différents selon des groupes d’USMs
		stats et graphiques associés




## Advanced use

You have learned so far how to use very useful but quite basic features of SticsRpacks packages.
Let's go now deeper into the packages functionalities.

### Modifying STICS files
SAMUEL-REMI 

Leur faire trouver les noms de fonction dans le REFERENCE du site web ... 
 
Montrer comment modifier les fichiers XML + faire graphe avant / après modification (e.g. code option type stress ou calcul ETP ou forcage).

Dire dans le see also que l'on peut aussi modifier les fichiers txt


### Generating STICS input Files

The different types of STICS input files can be generated using the [SticsRFiles package](https://sticsrpacks.github.io/SticsRFiles/index.html). 


#### Generate JavaStics Input files (XML format) 

The STICS XML files can be generated from CSV or Excel files.
except General parameters files and weather data files.


<div class="alert alert-info">
  See also: the vignette [Generating Stics XML files from tabulated data](https://sticsrpacks.github.io/SticsRFiles/articles/Generating_Stics_XML_files.html).
</div>


##### Exercises

#### Generate Stics Input files (Text format)

##### Exercises

#### Generating observation files 

Observation files can be generated from a data.frame using the function `gen_obs` of the [SticsRFiles package](https://sticsrpacks.github.io/SticsRFiles/index.html).

They are not used by the STICS model but can be used to plot comparison between observed and simulated variables and to optimize STICS parameters both in JavaStics and in [CroPlotR](https://sticsrpacks.github.io/CroPlotR/index.html) and [CroptimizR](https://sticsrpacks.github.io/CroptimizR/index.html) packages.
 
##### Exercises

**Pré-remplir un tableau de données avec des Dates en POSIXct et leur faire ajouter les colonnes ian, mo, jo + usm_name**
**Ou mieux: modifier gen_obs pour qu'il prenne en entrée la colonne Date plutot que ian, mo, jo ...)**
 
* Generate observation files for two usms named "usm1" and "usm2" including observations for a couple of dates for `lai(n)` and `mafruit` variables 

using the `gen_obs` function from a data.frame.

check with get_obs ?

hint
hint: `lai(n)` ????
tester avec lai_n
que se passe-t-il si on met n'importe quoi comme variable ?


```{r gen_obs, exercise=TRUE}
# Insert your R code here!
```

```{r get_param_xml1-hint-1}
# Soils are defined in the "sols.xml" file ...
```

```{r get_param_xml1-solution}
#SOLUTION:
get_param_xml(file.path(workspace_xml,"sols.xml"), c("argi"))
```

 
 
 
 
### Running simulations from R using stics_wrapper
REMI

#### The need of a wrapper

We saw earlier how to make a simulation using `run_javastics()`. This function simply call javastics from the command line to make simulations, which does the same as running an usm by opening the javastics user interface, but from R. 

You are probably familiar with the xml files used as inputs for javastics (*e.g.* `sols.xml` or `usms.xml`). These files list all the parameter values for one or several usms, whith potentially several soils, management practices and climate. Then when we run a usm, javastics reads the parameters values from the xml files, and writes new text files with only the values associated to this particular usm. Then it calls the stics model that reads those text files, make the simulation, and writes the outputs. 

So when we run a usm from javastics we have four different steps:

1. javastics reads the parameter values for the usm from the xml files
1. javastics writes text files with the parameter values that only describe this usm
1. javastics calls the stics model (*i.e* the stics executable)
1. the stics model reads the parameter values from the text files, make a simulation, and writes the output files (*i.e.* `mod_bXXX.sti` and `mod_sXXX.sti`)

This method is simpler for the user because it hides the technical details, but it has some downsides: it is slow because we have to create the text files each time we make a simulation, and we cannot parallelize the simulations because we share the same text files.

But no worries, SticsRPacks provides a solution to this!

Instead of creating the text files one usm at a time, we create them for all usms once into one folder for each, and then we directly call the stics model (*i.e* without using javastics) on these files. Using this method we are now able to run the simulations way faster, and in parallel! 

So how do we do that? Well, we use three steps:

- create the text files (`gen_usms_xml2txt()`)
- set-up our modeling environment (`stics_wrapper_options()`)
- run the simulations and import the results (`stics_wrapper()`)

#### Creating the text files

The first step is to create the text files for each usms in a separate directory. To do so, we use `gen_usms_xml2txt()`.

This tutorial stores a javastics installation in the `javastics` variable, and an example workspace in the `workspace_xml` variable.

Use `gen_usms_xml2txt()` with those variables to generate the text files for the usms "wheat", "banana" and "soybean":

```{r gen_usms_xml2txt, exercise=TRUE}
# Generate text files for usms "wheat" and "banana" and "soybean"
```

```{r gen_usms_xml2txt-hint-1}
# Use `javastics_path = javastics` and `workspace_path = workspace_xml`
```

```{r gen_usms_xml2txt-solution}
#SOLUTION:
gen_usms_xml2txt(javastics, workspace_xml, usms_list = c("wheat","banana","soybean"))
```

#### Setting-up the modeling environment

Before running the simulations using stics, we have to define some information about our modeling environment. This is done using `stics_wrapper_options()`. This function helps us define where is our installation of javastics, which stics executable to use (*e.g.* if you are using an executable from a research branch), which usms to simulate, etc...

Use `stics_wrapper_options()` to set-up our modeling environment using a parallel computation over the usms:

```{r stics_wrapper_options1, exercise=TRUE}
# Set-up the modeling environment
```

```{r stics_wrapper_options1-hint-1}
# Use `javastics_path = javastics` and `data_dir = workspace_xml`
```

```{r stics_wrapper_options1-solution}
#SOLUTION:
stics_wrapper_options(javastics_path = javastics, data_dir = workspace_xml, parallel = TRUE)
```

#### Running the simulations

Now everything is set-up to make our simulations! The model options we will use here are stored in the `wrapper_options` object with the options we just set up in the previous section. 

```{r prepare-wrapper, echo=FALSE}
wrapper_options = stics_wrapper_options(javastics_path = javastics, data_dir = workspace_xml, parallel = TRUE)
```

<!-- Here is its output: -->
<!-- ```{r wrapper_options_print, exercise.setup = "prepare-wrapper"} -->
<!-- wrapper_options -->
<!-- ``` -->

Now use `stics_wrapper()` to simulate the usm "wheat" only (not "banana" and "soybean") using the model options stored in `wrapper_options`, and return the leaf area index and the above-ground biomass:

```{r stics_wrapper1, exercise=TRUE, exercise.setup = "prepare-wrapper"}
# Simulate "wheat"
```

```{r stics_wrapper1-hint-1}
# The sit_names argument is used to choose the usms to simulate.
```

```{r stics_wrapper1-hint-2}
# Use get_var_info() to search for a variable name
```

```{r stics_wrapper1-solution, exercise.setup = "prepare-wrapper"}
#SOLUTION:
stics_wrapper(model_options = wrapper_options, sit_names = "wheat", var_names = c("lai(n)", "masec(n)"))
```

#### Parameter forcing

Sometimes it is useful to compare the outputs of the model with a different value for one or several parameters to better understand its impact or check if it is used. 

`stics_wrapper()` provides an easy way to do so with the `param_values` argument. This argument is given as a named vector.

> Make a simulation on the "wheat" usm and output just the lai, but this time with `dlaimax = 0.005` and `durvieF = 80`:

```{r stics_wrapper2, exercise=TRUE, exercise.setup = "prepare-wrapper"}
# Insert your R code here!
```

```{r stics_wrapper2-hint-1}
# `param_values` is given as a named vector, e.g. c("bdens" = 5, "hautmax" = 1.1)
```

```{r stics_wrapper2-solution, exercise.setup = "prepare-wrapper"}
#SOLUTION:
stics_wrapper(model_options = wrapper_options, sit_names = "wheat", var_names = "lai(n)", param_values = c("dlaimax" = 0.005, "durvieF" = 80))
```

#### Plot the comparison

Now we can plot the output of the two simulations using `plot()` and adding the simulation objects as arguments (optionally named). The results of both simulations are accessible in the `res1` and `res2` objects. The results of the simulations are stored in the `sim_list` field of each object. For example to get the results of the simulation for `res1`, we would do: `res1$sim_list`. 

```{r prepare-wrapper-plot, echo=FALSE}
wrapper_options = stics_wrapper_options(javastics_path = javastics, data_dir = workspace_xml, parallel = TRUE)

res1 = stics_wrapper(model_options = wrapper_options, sit_names = "wheat", var_names = c("lai_n", "masec_n"))
res2 = stics_wrapper(model_options = wrapper_options, sit_names = "wheat", var_names = "lai_n", param_values = c("dlaimax" = 0.005, "durvieF" = 80))
```

Then, to plot a simulation result, we would do:

```{r stics_wrapperplot, exercise=TRUE, exercise.setup = "prepare-wrapper-plot"}
# Execute the code below to create a plot of the results:
plot(res1$sim_list)
```

It is also possible to compare two simulations by listing them as arguments to the function one after the other.

> Plot the outputs of both simulations:

```{r stics_wrapper3, exercise=TRUE, exercise.setup = "prepare-wrapper-plot"}
# Insert your R code here!
```

```{r stics_wrapper3-hint-1}
# The simulation results can be listed one after the other in `plot()` for comparison
```

```{r stics_wrapper3-solution, exercise.setup = "prepare-wrapper-plot"}
#SOLUTION:
plot(res1$sim_list,res2$sim_list)
```

<div class="alert alert-info">
  Only the `lai` variable was requested on the second simulation, that's why there is only the first simulation for masec.
</div>

We can also name the simulations in our plot by naming them when passed as arguments to `plot()`.

> Plot the outputs of both simulations, and name the first simulation "original" and the second simulation "dlaimax=0.005, durvieF=80":

```{r stics_wrapper4, exercise=TRUE, exercise.setup = "prepare-wrapper-plot"}
# Insert your R code here!
```

```{r stics_wrapper4-hint-1}
# Name the arguments of the plot function, e.g. plot("sim1" = res1) would name the simulation "sim1" (but would not output it on the plot because there is only one)
```

```{r stics_wrapper4-solution, exercise.setup = "prepare-wrapper-plot"}
#SOLUTION:
plot("original" = res1$sim_list,"dlaimax=0.005, durvieF=80" = res2$sim_list)
```

### Estimating parameters
REMI

#### Workflow

Model calibration is an important part of most modeling work, but it can be complex and cumbersome. SticsRPacks provides a nice, standardized and easy way to estimate the values of parameters using observations: `estim_param()`.

The workflow for parameter estimation is made using several steps:

- Modeling environment set-up (`stics_wrapper_options()`). `estim_param()` uses `stics_wrapper()` under the hood to run each simulation, so we must set-up our modeling environment first as we saw before. 
- Import the observations for each usm (`get_obs()` + `filter_obs()`).
- Choose the parameter(s) to optimize and their boundary values.
- set-up options for the optimization method (see other section, not presented in this one for simplicity).
- run the parameter optimization (`optim_options`).

```{r prepare-param-estim, include=FALSE}
wrapper_options = stics_wrapper_options(javastics_path = javastics, data_dir = workspace_xml, parallel = TRUE)

obs = get_obs(workspace = workspace_xml, usm_name = "wheat")
obs_filt = filter_obs(obs, "lai_n", include = TRUE)
```

#### Modelling environment

We will simply re-use the modeling environment from the previous exercise: `wrapper_options`. This is how we did it:

```r
wrapper_options = stics_wrapper_options(javastics_path = javastics, data_dir = workspace_xml, parallel = TRUE)
```

#### Observations

> Import the observations for the "wheat" usm, and put it in an object called `obs`:

```{r estim_param_obs1, exercise=TRUE}
# Insert your R code here!
```

```{r estim_param_obs1-hint-1}
# Use `get_obs()`, workspace_xml is the workspace.
```

```{r estim_param_obs1-solution, exercise.setup = "prepare-param-estim"}
#SOLUTION:
obs = get_obs(workspace = workspace_xml, usm_name = "wheat")
```

> Filter the observations using `filter_obs()` for the lai only, and put it in an object called `obs_filt`:

```{r estim_param_obs2, exercise=TRUE}
# Insert your R code here!
```

```{r estim_param_obs2-hint-1}
# Use `get_obs()`, workspace_xml is the workspace.
```

```{r estim_param_obs2-solution, exercise.setup = "prepare-param-estim"}
#SOLUTION:
obs_filt = filter_obs(obs, "lai_n", include = TRUE)
```

#### Choose parameter

We have to choose which parameter will be optimized, and what are their respective boundaries values. To do so, we define a list of list for each parameter like so: 

```r
param_info = list()
param_info$dlaimax = list(lb=0.0005,ub=0.0025)
param_info$durvieF = list(lb=c(50,100),ub=c(400,450))
```

<div class="alert alert-danger">
  REMI: Continue here. 
  
1. Check if providing param_info like this above works (we will replace the other way to do it by this one already used also). 
1. And continue below
</div>
   
#### Modelling environment

> Now estimate the value of the `dlaimax` parameter using `estim_param()`, and put the results in an object called `optim_res`. You can use `wrapper_options` as the object for the argument `model_options`, and the `obs` object for the observations list. 

```{r estim_param2, exercise=TRUE, exercise.setup = "prepare-wrapper-plot"}
# Insert your R code here!
```

```{r estim_param2-hint-1}
# The simulation results can be listed one after the other in `plot()` for comparison
```

```{r estim_param2-solution, exercise.setup = "prepare-wrapper-plot"}
#SOLUTION:
optim_res = estim_param(obs_list = obs_filt,
                        model_function = stics_wrapper,
                        model_options = wrapper_options,
                        param_info = param_info)
```

estim_param
	méthode par défaut
	Afficher résultats produits et plots

plots avec CroPlotR : run du wrapper sur les résultats de l'optimisation puis plot sim-obs

### Succession culturales
REMI
 
	stics_wrapper
	CroPlotR : dynamique sur succession
estim_param

### Intercrop
SAMUEL

	
	get_sim
CroPlotR : plot + stats ?	
 
### Advanced parameter estimation 
REMI

plus avancé :
	contraintes inégalités
	paramètres par groupe (varietal / spécifique)
méthode bayesienne
transfo des sim et obs ???

